<a href="https://colab.research.google.com/github/Komuravelli-Kaustubh/AI-Agent-FinancialAgent/blob/main/Agent-order_status_assistant.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Kaustubh"
COLLABORATORS = "ksb"
```

---

```python
# '''
# return-policy.txt
# -----------------
# return_policy:
# Our return policy allows returns within 30 days of purchase.
# for Tablet , the return policy allows return within 10 days of purchase.

# order_tracking:
# You can track your order status using your order ID.

# product_info:
# Product information includes specifications and features.

# orders.json
# ------------
# {
#     "123": {
#         "customer_name": "John Doe",
#         "order_date": "2022-01-01",
#         "products": ["Laptop", "Monitor"],
#         "total": 1500.00
#     },
#     "124": {
#         "customer_name": "Jane Doe",
#         "order_date": "2022-01-15",
#         "products": ["Phone", "Keyboard"],
#         "total": 800.00
#     },
#     "125": {
#         "customer_name": "Bob Smith",
#         "order_date": "2022-02-01",
#         "products": ["Tablet", "Laptop"],
#         "total": 2500.00
#     }
# }


# product_info.json
# -----------------
# {
#     "Laptop": "A high-performance laptop with 16GB RAM and 512GB SSD.",
#     "Phone": "A smartphone with a 6.1-inch display and 128GB storage.",
#     "Tablet": "A tablet with a 10.5-inch display and 256GB storage.",
#     "Monitor": "A 24-inch full HD monitor.",
#     "Keyboard": "A mechanical keyboard with RGB backlighting."
# }


# order_status.json
# -----------------
# {
#     "123": "Delivered",
#     "124": "Delivered",
#     "125": "Delivered",
# }

# '''
```

```python
# ✅ Create return-policy.txt
with open("return-policy.txt", "w") as f:
    f.write("""return_policy:
Our return policy allows returns within 30 days of purchase.
for Tablet , the return policy allows return within 10 days of purchase.

order_tracking:
You can track your order status using your order ID.

product_info:
Product information includes specifications and features.
""")

# ✅ Create orders.json
import json

orders = {
    "123": {
        "customer_name": "John Doe",
        "order_date": "2022-01-01",
        "products": ["Laptop", "Monitor"],
        "total": 1500.00
    },
    "124": {
        "customer_name": "Jane Doe",
        "order_date": "2022-01-15",
        "products": ["Phone", "Keyboard"],
        "total": 800.00
    },
    "125": {
        "customer_name": "Bob Smith",
        "order_date": "2022-02-01",
        "products": ["Tablet", "Laptop"],
        "total": 2500.00
    }
}

with open("orders.json", "w") as f:
    json.dump(orders, f, indent=4)

# ✅ Create product_info.json
product_info = {
    "Laptop": "A high-performance laptop with 16GB RAM and 512GB SSD.",
    "Phone": "A smartphone with a 6.1-inch display and 128GB storage.",
    "Tablet": "A tablet with a 10.5-inch display and 256GB storage.",
    "Monitor": "A 24-inch full HD monitor.",
    "Keyboard": "A mechanical keyboard with RGB backlighting."
}

with open("product_info.json", "w") as f:
    json.dump(product_info, f, indent=4)

# ✅ Create order_status.json
order_status = {
    "123": "Delivered",
    "124": "Delivered",
    "125": "Delivered"
}

with open("order_status.json", "w") as f:
    json.dump(order_status, f, indent=4)

```

```python
from langchain.tools import Tool
import json
# from langchain_openai import ChatOpenAI
import os
from datetime import date

def get_order_status_tool(order_id):
    if not order_id.startswith("order_id = "):
        return "Please provide an order ID in the format: 'order_id = <your_order_id>'"

    try:
        order_id = order_id.replace("order_id = ", "")
        with open('order_status.json', 'r') as file:
            order_status = json.load(file)
        return order_status.get(str(order_id), "Order ID not found")
    except Exception as e:
        print(f"Error reading order_status.json: {e}")
        return "Order ID not found"

def save_order_status_returned_tool(order_id):
    # Check for the correct format of order_id
    if not order_id.startswith("order_id = "):
        return "Please provide an order ID in the format: 'order_id = <your_order_id>'"

    try:
        # Read the existing data from the file
        try:
            with open('order_status.json', 'r') as file:
                order_status = json.load(file)
        except FileNotFoundError:
            order_status = {}
        except json.JSONDecodeError:
            order_status = {}

        # Extract the actual order ID
        order_id = order_id.replace("order_id = ", "")

        # Check if the order ID exists in the current data
        if order_id not in order_status:
            return "Order ID not found"

        # Update the specific order status to "Returned"
        order_status[str(order_id)] = "Returned"

        # Write the updated data back to the file
        with open('order_status.json', 'w') as file:
            json.dump(order_status, file, indent=4)

        return "Order status updated successfully."

    except Exception as e:
        print(f"Error saving order_status.json: {e}")
        return "Failed to save order status."

def get_product_info_tool(product_name):
    if not product_name.startswith("product_name = "):
        return "Please provide a product name in the format: 'product_name = <your_product_name>'"

    try:
        product_name = product_name.replace("product_name = ", "")
        with open('product_info.json', 'r') as file:
            product_info = json.load(file)
        return product_info.get(product_name, "Product not found")
    except Exception as e:
        print(f"Error reading product_info.json: {e}")
        return "Product not found"

def get_order_info_tool(order_id):
    if not order_id.startswith("order_id = "):
        return "Please provide an order ID in the format: 'order_id = <your_order_id>'"

    try:
        order_id = order_id.replace("order_id = ", "")
        with open('orders.json', 'r') as file:
            orders = json.load(file)
        return orders.get(str(order_id), "Order ID not found")
    except Exception as e:
        print(f"Error reading orders.json: {e}")
        return "Order ID not found"

def get_today_date(today=date.today()):
    return "today is {}".format(today)


input_tool = Tool(
    name="get_input_from_user",
    func=input,
    description="Get the input from the user."
)

order_status_tool = Tool(
    name="get_order_status",
    func=get_order_status_tool,
    description="Get the status of an order given an order ID."
)

product_info_tool = Tool(
    name="get_product_info",
    func=get_product_info_tool,
    description="Get information about a product given its name."
)

order_info_tool = Tool(
    name="get_order_info",
    func=get_order_info_tool,
    description="Get information about an order given its ID."
)

order_status_returned_tool = Tool(
    name="save_order_status_returned",
    func=save_order_status_returned_tool,
    description="Save the status of an order given an order ID."
)

get_today_date_tool = Tool(
    name="get_today_date",
    func=get_today_date,
    description="Get the current day date."
)

```

<!-- grade_id: cell-561005cecbf5c5c1, points: 50, grade: true, solution: false -->
```python
order_status = order_status_tool.invoke("order_id = 124")
assert(order_status=='Delivered')
```

```python
# os.environ["OPENAI_API_KEY"] = "<place your open ai api key here>"
# llm = ChatOpenAI(model="gpt-4o", temperature=0)
# 🛠️ Install required packages
!pip install -q langchain-groq groq python-dotenv

# ✅ Set API key directly in Colab (you can paste your Groq key here)
import os
os.environ["GROQ_API_KEY"] = "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"  # 🔒 Replace with your actual key

# ✅ Initialize Groq LLM (using LangChain)
from langchain_groq import ChatGroq

llm = ChatGroq(
    model_name="llama3-70b-8192",  # Options: mixtral-8x7b-32768, gemma-7b-it
    temperature=0
)
```

Output:
```
[?25l   [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m0.0/126.7 kB[0m [31m?[0m eta [36m-:--:--[0m[2K   [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m126.7/126.7 kB[0m [31m4.3 MB/s[0m eta [36m0:00:00[0m
[?25h
```

```python
!pip install langchain-community
```

Output:
```
Collecting langchain-community
  Downloading langchain_community-0.3.21-py3-none-any.whl.metadata (2.4 kB)
Requirement already satisfied: langchain-core<1.0.0,>=0.3.51 in /usr/local/lib/python3.11/dist-packages (from langchain-community) (0.3.51)
Requirement already satisfied: langchain<1.0.0,>=0.3.23 in /usr/local/lib/python3.11/dist-packages (from langchain-community) (0.3.23)
Requirement already satisfied: SQLAlchemy<3,>=1.4 in /usr/local/lib/python3.11/dist-packages (from langchain-community) (2.0.40)
Requirement already satisfied: requests<3,>=2 in /usr/local/lib/python3.11/dist-packages (from langchain-community) (2.32.3)
Requirement already satisfied: PyYAML>=5.3 in /usr/local/lib/python3.11/dist-packages (from langchain-community) (6.0.2)
Requirement already satisfied: aiohttp<4.0.0,>=3.8.3 in /usr/local/lib/python3.11/dist-packages (from langchain-community) (3.11.15)
Requirement already satisfied: tenacity!=8.4.0,<10,>=8.1.0 in /usr/local/lib/python3.11/dist-packages (from la
```

```python
!pip install faiss-cpu
```

Output:
```
Collecting faiss-cpu
  Downloading faiss_cpu-1.10.0-cp311-cp311-manylinux_2_28_x86_64.whl.metadata (4.4 kB)
Requirement already satisfied: numpy<3.0,>=1.25.0 in /usr/local/lib/python3.11/dist-packages (from faiss-cpu) (2.0.2)
Requirement already satisfied: packaging in /usr/local/lib/python3.11/dist-packages (from faiss-cpu) (24.2)
Downloading faiss_cpu-1.10.0-cp311-cp311-manylinux_2_28_x86_64.whl (30.7 MB)
[2K   [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m30.7/30.7 MB[0m [31m56.7 MB/s[0m eta [36m0:00:00[0m
[?25hInstalling collected packages: faiss-cpu
Successfully installed faiss-cpu-1.10.0
```

```python
from langchain.vectorstores import FAISS
# from langchain.embeddings.openai import OpenAIEmbeddings
from langchain.embeddings import HuggingFaceEmbeddings
from langchain.text_splitter import RecursiveCharacterTextSplitter
from langchain.document_loaders import TextLoader
from langchain.llms import OpenAI
from langchain.document_loaders import PyMuPDFLoader

# Load and index documents
loader = TextLoader("return-policy.txt")
documents = loader.load()

# Split documents into chunks
text_splitter = RecursiveCharacterTextSplitter(chunk_size=1000, chunk_overlap=20)
split_docs = text_splitter.split_documents(documents)

# Create embeddings and build FAISS index
# embeddings = OpenAIEmbeddings()
embeddings = HuggingFaceEmbeddings(model_name="sentence-transformers/all-MiniLM-L6-v2")
vector_store = FAISS.from_documents(split_docs, embeddings)
retriever = vector_store.as_retriever(search_type="similarity_score_threshold", search_kwargs={"score_threshold": 0.2})

from langchain.tools.retriever import create_retriever_tool
from langchain.agents import AgentExecutor, load_tools

tool_search = create_retriever_tool(
    retriever=retriever,
    name="search_return_policy",
    description="Searches and returns policy, product, etc."
)

```

```python
tools = load_tools([], llm=llm)
tools.append(tool_search)
tools.append(input_tool)
tools.append(order_status_tool)
tools.append(product_info_tool)
tools.append(order_info_tool)
tools.append(order_status_returned_tool)
tools.append(get_today_date_tool)

from langchain import hub
prompt = hub.pull("hwchase17/react")

# only creates the logical steps for us
from langchain.agents import AgentExecutor, create_react_agent

react_agent = create_react_agent(llm, tools, prompt)
agent_executor = AgentExecutor(
    agent=react_agent, tools=tools, verbose=True, handle_parsing_errors=True
)

def prompt_user_for_info(query):
    response = agent_executor.invoke({"input": query})
    if "Please provide" in response["output"]:
        new_input = input(response["output"])  # Prompt the user for additional input
        return prompt_user_for_info(new_input)
    else:
        return response

```

Output:
```
/usr/local/lib/python3.11/dist-packages/langsmith/client.py:278: LangSmithMissingAPIKeyWarning: API key must be provided when using hosted LangSmith API
  warnings.warn(
```

<!-- grade_id: cell-5475e924e91a1316, points: 50, grade: true, solution: false -->
```python
query = "what is the order status of order_id 125?"
final_response = prompt_user_for_info(query)
print(final_response)
assert("Delivered" in final_response['output'])
```

Output:
```
[1m> Entering new AgentExecutor chain...[0m
[32;1m[1;3mThought: I need to get the order status of order_id 125.

Action: get_order_status
Action Input: 125[0m[38;5;200m[1;3mPlease provide an order ID in the format: 'order_id = <your_order_id>'[0m[32;1m[1;3mIt seems like I need to provide the order ID in a specific format.

Action: get_order_status
Action Input: 'order_id = 125'[0m[38;5;200m[1;3mPlease provide an order ID in the format: 'order_id = <your_order_id>'[0m[32;1m[1;3mIt seems like I'm stuck in an infinite loop! Let me try again.

Action: get_order_status
Action Input: order_id = 125[0m[38;5;200m[1;3mDelivered[0m[32;1m[1;3mI now know the final answer
Final Answer: The order status of order_id 125 is Delivered.[0m

[1m> Finished chain.[0m
{'input': 'what is the order status of order_id 125?', 'output': 'The order status of order_id 125 is Delivered.'}
```

```python
query = "get product information of the order?"
final_response = prompt_user_for_info(query)
print(final_response)
```

Output:
```
[1m> Entering new AgentExecutor chain...[0m
[32;1m[1;3mThought: I need to get the order ID from the user to get the product information.

Action: get_input_from_user
Action Input: "Please enter the order ID: "[0mPlease enter the order ID: 123
[33;1m[1;3m123[0m[32;1m[1;3mThought: Now that I have the order ID, I need to get the order information using the order ID.

Action: get_order_info
Action Input: 123[0m[33;1m[1;3mPlease provide an order ID in the format: 'order_id = <your_order_id>'[0m[32;1m[1;3mIt seems like the observation from the `get_order_info` action is not providing the actual order information, but rather a prompt to provide the order ID in a specific format. Let's try again.

Action: get_input_from_user
Action Input: "Please enter the order ID in the format 'order_id = <your_order_id>': "[0mPlease enter the order ID in the format 'order_id = <your_order_id>': order_id = 123
[33;1m[1;3morder_id = 123[0m[32;1m[1;3mThought: Now that I have the order ID 
```

```python
query = "Return my product which was ordered?"
final_response = prompt_user_for_info(query)
print(final_response)
```

Output:
```
[1m> Entering new AgentExecutor chain...[0m
[32;1m[1;3mThought: I need to know the order ID to return the product.

Action: get_input_from_user
Action Input: "Please enter your order ID: "[0mPlease enter your order ID: order_id = 125
[33;1m[1;3morder_id = 125[0m[32;1m[1;3mThought: Now that I have the order ID, I need to check the status of the order.

Action: get_order_status
Action Input: order_id = 125[0m[38;5;200m[1;3mDelivered[0mWARNING:langchain_core.vectorstores.base:No relevant docs were retrieved using the relevance score threshold 0.2
[32;1m[1;3mThought: Since the order status is "Delivered", I need to check the return policy for this product.

Action: search_return_policy
Action Input: query = "return policy"[0m[36;1m[1;3m[0mWARNING:langchain_core.vectorstores.base:No relevant docs were retrieved using the relevance score threshold 0.2
[32;1m[1;3mThought: I need to extract the relevant information from the return policy to determine if the product can be
```

```python

```